# Bisseção em 2D

Imagine que temos duas funções $f$ e $g$ definidas de $\mathbb R^2$ em $\mathbb R$.
(ou seja, um sistema - não linear - com duas equações e duas incógnitas)

Como encontrar uma raiz deste sistema?

```python
%pylab inline
```

# Um "quadrado" para fazer bisseção

Queremos achar duas curvas $a(t)$ e $b(t)$ tais que

- $f(a(t)) > 0$
- $f(b(t)) < 0$

Assim, poderemos construir pontos $c(t)$ tais que $f(c(t)) \sim 0$
por bisseção no segmento de reta $\overline{a(t) \ b(t)}$ (ou qualquer outra curva ligando estes pontos).
Então, "basta" encontrar dois valores de $t$ onde $g(c(t))$ tenha sinais opostos para encontrar uma raiz do sistema
$f = 0 \ , \ g = 0$.

Como você vai ver, o mais difícil deste método é encontrar as curvas $a(t)$ e $b(t)$ satisfazendo todas estas hipóteses.
Depois, as aplicações da bisseção para achar as raízes duplas são relativamente fáceis.

## Primeira parte: encontrando as curvas $a(t)$ e $b(t)$.

Precisamos (da mesma forma que na bisseção em $\mathbb R$) de pontos iniciais que satisfaçam algumas condições.

Uma ideia natural é construir uma "espiral" de pontos e testá-los, e depois construir as curvas $a(t)$ e $b(t)$
"ligando" estes pontos.

Comecemos pela espiral.

### Questão 1: Espiral

Escreva uma função `espiral(n)` que constrói $n$ pontos de norma cada vez maior, partindo de $(0,0)$ e girando à esquerda.

Ela deve retornar os pontos percorridos em uma lista.

<!-- grade_id: def_espiral, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

E agora, os asserts:

<!-- grade_id: espiral_1, points: 2, grade: true, solution: false -->
```python
# Tamanho
assert(len(espiral(30)) == 30)
assert(len(espiral(100)) == 100)
```

<!-- grade_id: espiral_2, points: 3, grade: true, solution: false -->
```python
# Norma maior
e = espiral(100)
for x,y in zip(e[:-1], e[1:]):
    assert(norm(x) < norm(y))
```

<!-- grade_id: espiral_3, points: 3, grade: true, solution: false -->
```python
# Girando à esquerda
e = espiral(100)
for x,y in zip(e[1:-1], e[2:]):
    assert(cross(x,y) > 0)
```

### Questão 2: Pontos com $f$ de mesmo sinal

Queremos achar uma curva onde $f$ será positiva, e outro onde $f$ será negativa.

Escreva uma função que, dada uma lista de pontos, retorna duas listas:
- uma com os pares de pontos consecutivos onde $f$ é positiva,
- a outra com os pares onde $f$ é negativa.

Considere que a função `f` recebe apenas um argumento (**uma** lista, **um** `array`, **um** par de pontos, ...).
Desta forma, a sua função `consecutivos` funcionará em qualquer dimensão!

Um exemplo simples de $f$ e seu código:

```python
def f(p):
    x,y = p
    return x**2 + 3*x - 20*y + y**2
```

<!-- grade_id: def_consecutivos, points: 0, grade: false, solution: true -->
```python
def consecutivos(f,pts):
    # Construa essas listas com append ;-)
    l_pos = []
    l_neg = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return l_pos, l_neg
```

<!-- grade_id: consec_1, points: 2, grade: true, solution: false -->
```python
pts = [(0, 0),
       (1, 0),
       (1, 1),
       (-1, 1),
       (-1, -1),
       (2, -1),
       (2, 2),
       (-2, 2),
       (-2, -2),
       (3, -2),
       (3, 3)]
assert( consecutivos(f, pts) ==
        ([((-1, -1), (2, -1)), ((-2, -2), (3, -2))],
         [((1, 1), (-1, 1)), ((2, 2), (-2, 2))])
      )
```

Encontre agora uma espiral suficientemente grande com pontos de sinal positivo e negativo:

<!-- grade_id: n_consec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: consec_2, points: 2, grade: true, solution: false -->
```python
l_pos, l_neg = consecutivos(f, espiral(n_grande_espiral))
assert( len(l_pos) > 0 )
assert( len(l_neg) > 0 )
```

### Questão 3: Segmentos (ou "CVGA strikes back")

Agora, faça uma função `segmento(p,q)` recebendo dois pontos e que retorna **uma função** `s(t)` que representa o segmento entre eles.
(Poderíamos representar o segmento "simplesmente" por seus dois pontos extremos, e daí não haveria nada a fazer nesta questão. A vantagem de representar o segmento por uma função é que podemos acessar todos pontos do segmento de forma mais fácil e homogênea)

Esta função será útil no mínimo duas vezes:
- Para construir as curvas $a(t)$ e $b(t)$, que serão segmentos (para simplificar a nossa vida) entre pontos consecutivos de mesmo sinal de $f$,
- Para construir os segmentos $\overline{a(t) \ b(t)}$

<!-- grade_id: def_segmento, points: 0, grade: false, solution: true -->
```python
def segmento(p,q):
    # Vetorizando, vai ficar mais fácil fazer as contas!
    p = array(p)
    q = array(q)
    def s(t):
        # YOUR CODE HERE
        raise NotImplementedError()
    return s
```

Um exemplo para ver como funciona (e também para mostrar que esta função, _na verdade_ faz a reta entre eles se passarmos $t$ fora do intervalo $[0,1]$.)

```python
s = segmento((0,1), (1,0))
s(0), s(.5), s(1), s(2)
```

<!-- grade_id: segmento_1, points: 2, grade: true, solution: false -->
```python
s = segmento((-3,-1), (2,3))
assert(allclose( s(0.6), [0,1.4], rtol=1e-14, atol=1e-14 ))
```

<!-- grade_id: segmento_2, points: 2, grade: true, solution: false -->
```python
s = segmento((2,3), (4,-9))
assert(all( s(.125) == [2.25, 1.5] ))
```

Testando que as extremidades inicial e final estão no lugar certo:

<!-- grade_id: segmento_3, points: 3, grade: true, solution: false -->
```python
ps = randn(5,2)
qs = randn(5,2)
for p,q in zip(ps,qs):
    s = segmento(p,q)
    assert(all( s(0) == p ))
    assert(all( s(1) == q ))
```

### Questão 4: Segmentos totalmente positivos e negativos de $f$

Use um gráfico para mostrar que, em algum dos segmentos encontrados pela função `consecutivos`,
a função `f` é sempre positiva.
Faça depois o mesmo para achar um segmento totalmente negativo.

```python
l_pos, l_neg = consecutivos(f, espiral(10))
```

<!-- grade_id: seg_graph, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: seg_graph_2, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Segunda parte: Achando zeros de $f$

Agora, para cada valor de $t$, queremos achar um zero de $f$ no segmento $\overline{a(t)\ b(t)}$.
Mas lembre que a bisseção só sabe dividir um intervalo **real** por dois, não um segmento qualquer.

Assim, vamos ter que criar, **para cada $t$**, uma nova função $f_t(r)$, que é a restrição de $f$ à reta $\overline{a(t)\ b(t)}$.

### Questão 5: Segmentos de segmentos

Escreva uma função `segmento_t(a,b,t)` que cria o segmento entre os pontos $a(t)$ e $b(t)$.

<!-- grade_id: def_seg_t, points: 0, grade: false, solution: true -->
```python
def segmento_t(a,b,t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: seg_t_1, points: 2, grade: true, solution: false -->
```python
a = segmento((0,0), (1,0))
b = segmento((0,1), (1,1))
st = segmento_t(a,b,0.25)
assert(all( st(0.4) == [0.25,0.4] ))
```

<!-- grade_id: seg_t_2, points: 3, grade: true, solution: false -->
```python
for t in randn(10):
    st = segmento_t(a,b,t)
    r = randn(1)[0]
    assert(allclose( st(r), [t,r], rtol=1e-14, atol=1e-14 )), "{} {} {}".format(st(r), [t,r], st(r) - [t,r])
```

### Questão 6: Desenhando segmentos

Escreva uma função que, dado um segmento (ou seja, uma função $s(t)$!), desenha o segmento de reta associado.

<!-- grade_id: def_plota_seg, points: 0, grade: false, solution: true -->
```python
def plota_segmento(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: plota_seg_1, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Questão 7: Vendo a homotopia

(Não se assuste se você nunca ouviu falar de homotopia, ou não lembra o que é :-) . Você vai ver)

Os segmentos $a$ e $b$ são os lados "de baixo" e "de cima" de um quadrado.
Os segmentos **entre os pontos** $a(t)$ e $b(t)$ são, portanto, segmentos verticais "dentro do quadrado".

```python
plota_segmento(a)
plota_segmento(b)
for t in linspace(0,1,6):
    plota_segmento(segmento_t(a,b,t))
axis([-1,2,-1,2]);
```

Crie os segmentos dos lados "esquerdo" e "direito" do quadrado, e desenhe a grade "transposta".

<!-- grade_id: transposta, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Terceira parte: A curva $c(t)$

Vejamos o que já construímos:

- Dois segmentos $a$ e $b$ ;
- em $a$ a função $f$ é positiva, em $b$ ela é negativa ;
- para cada $t$, um segmento $s_t$ ligando os pontos $a(t)$ (com o outro lado $b$.

Assim, para cada $t$, podemos achar um ponto $c(t)$, no segmento $s_t$, por bisseção,
onde $f(c(t)) = 0$ (na verdade, será apenas _próximo_ de zero).

```python
def bissecao(f, a, b, prec=1e-8):
    """ Essa versão da bisseção garante que você não vai passar valores inválidos dos pontos extremos. """
    assert(f(a) * f(b) <= 0)
    
    def dividir(x,y):
        z = (x+y)/2
        err = max(abs(z-x), abs(z-y))
        if err < prec:
            return z

        if (f(x)*f(z) <= 0):
            return dividir(x,z)
        else:
            return dividir(z,y)

    return dividir(a,b)
```

```python
def construir_c_de_t(f,a,b):
    def c(t, prec=1e-8):
        # Para cada  t:
        # - constrói o segmento  s_t,
        # - restringe a função  f  a este segmento,
        # - faz a bisseção
        # - retorna o ponto (do segmento) correspondente
        s_t = segmento(a(t), b(t))
        def f_t(r):
            return f(s_t(r))
        r_zero = bissecao(f_t, 0,1, prec)
        return s_t(r_zero)
    return c
```

### Questão 8: Gráficos de uma lista de pontos

Se construirmos a curva $c(t)$ e depois calcularmos alguns valores, vamos ter uma lista de pontos.
Para fazer uma curva, temos que ter uma lista de coordenadas x, e depois uma lista de coordenadas y.
Escreva uma função `linha(pts)` que recebe uma lista de pontos e faz o gráfico.

<!-- grade_id: def_linha, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Em particular, `linha([p,q])` desenha o segmento $\overline{p \, q}$ ;-)

<!-- grade_id: graf_linha_1, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: graf_linha_2, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Questão 9: Traçando a curva $c(t)$

Faça um gráfico da curva $c(t)$ usando os dois segmentos "totalmente positivos" e "totalmente negativos" que você achou na questão 4.

```python
s_plus = s
s_minus = s_minus
```

<!-- grade_id: pts_c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: tracar_c, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Talvez você tenha escolhido a orientação "errada" para um dos segmentos (e em vez de ter um "quadrado", você tem um "X").
Talvez você tenha escolhido a orientação "certa".

Construa o outro caso, ligando os segmentos "ao contrário".
A função seguinte troca a orientação de um segmento pra você.

```python
def s_minus_rev(t):
    return s_minus(1-t)
```

```python
c = construir_c_de_t(f, s, s_minus_rev)
ts = arange(0,1,0.1)
ct = [c(t) for t in ts]
```

<!-- grade_id: c_reversa, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Quarta parte: Achando um zero de $g$ também

Até aqui, _nada_ depende de $g$.
Mas agora vamos pedir que $g$ **troque de sinal** na curva $c(t)$.

```python
def g(p):
    x,y = p
    return -cos(x) - sin(y**2) + x*y
```

Verificando que realmente troca de sinal:

```python
g(c(0)), g(c(1))
```

Agora, podemos fazer bisseção na função $g \circ c$!

### Questão 10: Escreva a função

<!-- grade_id: def_g_bola_c, points: 0, grade: false, solution: true -->
```python
def g_bola_c(r):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: g_bola_c_1, points: 3, grade: true, solution: false -->
```python
for r in arange(0,1,0.1):
    assert (g_bola_c(r) == g(c(r)))
```

### Questão 11: A raiz comum

Encontre uma raiz de $g \circ c$.
Como na curva $c(t)$ os pontos são raizes de $f$, você terá encontrado uma raiz comum!

<!-- grade_id: raiz_comum, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
f(raiz_comum), g(raiz_comum)
```

<!-- grade_id: v_raiz_comum, points: 2, grade: true, solution: false -->
```python
assert( abs(f(raiz_comum)) < 1e-7 )
assert( abs(g(raiz_comum)) < 1e-7 )
```

### Questão 12: Mais precisão

Agora, modifique alguma coisa acima para encontrar uma raiz mais precisa de $(f,g)$.

<!-- grade_id: nova_g, points: 2, grade: true, solution: true -->
```python
def g_bola_c(r):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: encontrando_precisao, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
f(raiz_comum), g(raiz_comum)
```